<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<!-- question: 1 -->
  <question type="stack">
    <name><text>PDE Basics - Q1: IC vs BC (Time)</text></name>
    <questiontext format="html">
      <text><![CDATA[<p>Determine whether the following condition is an <strong>initial condition</strong> or a <strong>boundary condition</strong>:</p>
<p>\[ u(x, {@t0@}) = {@expr@} \]</p>
<p>The condition is: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>The condition is given at \(t = {@t0@}\), which is the starting time. Therefore, it is an <strong>initial condition</strong>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2025073100</text></stackversion>
    <questionvariables>
      <text>t0: 0;
k: rand(5)+1;
expr: x + k;
ta1: [[1, true, "Initial Condition"], [2, false, "Boundary Condition"]];</text>
    </questionvariables>
    <specificfeedback format="html"><text></text></specificfeedback>
    <questionnote format="html"><text><![CDATA[<p>u(x,{@t0@})={@expr@}. IC.</p>]]></text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta1)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
    </prt>
  </question>

<!-- question: 2 -->
  <question type="stack">
    <name><text>PDE Basics - Q2: IC vs BC (Space)</text></name>
    <questiontext format="html">
      <text><![CDATA[<p>Determine whether the following condition is an <strong>initial condition</strong> or a <strong>boundary condition</strong>:</p>
<p>\[ u({@x0@}, t) = {@val@} \]</p>
<p>The condition is: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>The condition is specified at \(x = {@x0@}\), a fixed point in the spatial domain. Therefore, it is a <strong>boundary condition</strong>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2025073100</text></stackversion>
    <questionvariables>
      <text>x0: rand(5)+1;
val: rand(10);
ta1: [[1, false, "Initial Condition"], [2, true, "Boundary Condition"]];</text>
    </questionvariables>
    <specificfeedback format="html"><text></text></specificfeedback>
    <questionnote format="html"><text><![CDATA[<p>u({@x0@},t)={@val@}. BC.</p>]]></text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta1)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
    </prt>
  </question>

<!-- question: 3 -->
  <question type="stack">
    <name><text>PDE Basics - Q3: Simple PDE Solve</text></name>
    <questiontext format="html">
      <text><![CDATA[<p>Solve the PDE:</p>
<p>\[ \frac{\partial u}{\partial t} = 0 \]</p>
<p>subject to the initial condition \( u(x, 0) = {@ic_expr@} \).</p>
<p>\( u(x, t) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Since \(\frac{\partial u}{\partial t} = 0\), \(u\) does not depend on \(t\), so \(u(x,t) = f(x)\).</p>
<p>Applying the initial condition: \(f(x) = {@ic_expr@}\).</p>
<p>Therefore \(u(x,t) = {@ic_expr@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2025073100</text></stackversion>
    <questionvariables>
      <text>a: rand(5)+1;
b: rand(5)+1;
ic_expr: a*x^2 + b;
ta1: ic_expr;</text>
    </questionvariables>
    <specificfeedback format="html"><text></text></specificfeedback>
    <questionnote format="html"><text><![CDATA[<p>u(x,t)={@ta1@}.</p>]]></text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
    </prt>
  </question>

<!-- question: 4 -->
  <question type="stack">
    <name><text>PDE Basics - Q4: Dirichlet Type</text></name>
    <questiontext format="html">
      <text><![CDATA[<p>Identify the type of boundary condition:</p>
<p>\[ u({@x0@}, t) = {@val@} \]</p>
<p>This is a: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>The value of \(u\) is specified directly at the boundary. Therefore, this is a <strong>Dirichlet boundary condition</strong>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2025073100</text></stackversion>
    <questionvariables>
      <text>x0: 0;
val: rand(10)+1;
ta1: [[1, true, "Dirichlet"], [2, false, "Neumann"], [3, false, "Robin"]];</text>
    </questionvariables>
    <specificfeedback format="html"><text></text></specificfeedback>
    <questionnote format="html"><text><![CDATA[<p>u({@x0@},t)={@val@}. Dirichlet.</p>]]></text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta1)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
    </prt>
  </question>

<!-- question: 5 -->
  <question type="stack">
    <name><text>PDE Basics - Q5: Neumann Type</text></name>
    <questiontext format="html">
      <text><![CDATA[<p>Identify the type of boundary condition:</p>
<p>\[ \frac{\partial u}{\partial x}({@x0@}, t) = {@val@} \]</p>
<p>This is a: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>The derivative of \(u\) is specified at the boundary. Therefore, this is a <strong>Neumann boundary condition</strong>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2025073100</text></stackversion>
    <questionvariables>
      <text>x0: 1;
val: 0;
ta1: [[1, false, "Dirichlet"], [2, true, "Neumann"], [3, false, "Robin"]];</text>
    </questionvariables>
    <specificfeedback format="html"><text></text></specificfeedback>
    <questionnote format="html"><text><![CDATA[<p>u_x({@x0@},t)={@val@}. Neumann.</p>]]></text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta1)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
    </prt>
  </question>

</quiz>
