<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 169560 -->
  <question type="stack">
    <name>
      <text>Complex Analysis - Cauchy Integral Formula</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Using Cauchy's Integral Formula, evaluate the following integral:</p>
<p>\[ \oint_{|z|={@R@}} \frac{e^{ {@n@} z}}{z - {@a@}} dz \]</p>
<p>Answer: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>The integral is of the form \(\oint_C \frac{f(z)}{z-z_0} dz\) with:</p>
<ul>
    <li>\(f(z) = e^{ {@n@} z}\)</li>
    <li>\(z_0 = {@a@}\)</li>
    <li>\(C: |z| = {@R@}\)</li>
</ul>
<p>Since \(e^{ {@n@} z}\) is entire and \(|{@a@}| = {@a@} < {@R@}\), the point \(z_0 = {@a@}\) is inside the circle \(|z| = {@R@}\).</p>
<p>By Cauchy's Integral Formula:</p>
<p>\[ \oint_{|z|={@R@}} \frac{e^{ {@n@} z}}{z - {@a@}} dz = 2\pi i f(z_0) = 2\pi i \cdot e^{ {@n@} \cdot {@a@} } = 2\pi i e^{ {@n*a@} } \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>n:rand(3)+1;    /* coefficient in exponent */
a:rand(3)+1;    /* pole location */
R:a+1;          /* radius, ensures pole is inside */
ta1:2*%pi*%i*%e^(n*a);
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Integral of e^({@n@}z)/(z-{@a@}) over |z|={@R@}. Answer: 2*pi*i*e^{@n*a@}.</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>
