<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<question type="stack">
  <name><text>PDE BVP - Q1: u_xx=0, u(0)=a, u(L)=b</text></name>
  <questiontext format="html"><text><![CDATA[<p>Solve the boundary value problem:</p>
<p>\[\frac{\partial^2 u}{\partial x^2} = 0, \quad u(0) = {@u0@}, \quad u({@L@}) = {@uL@}.\]</p>
<p>\(u(x) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text></questiontext>
  <generalfeedback format="html"><text><![CDATA[<h3>Solution</h3>
<p>Integrate twice: \(u(x) = C_1 x + C_2\).</p>
<p>From \(u(0) = {@u0@}\): \(C_2 = {@u0@}\).</p>
<p>From \(u({@L@}) = {@uL@}\): \({@L@}C_1 + {@u0@} = {@uL@}\), so \(C_1 = {@C1@}\).</p>
<p>Therefore \(u(x) = {@ta1@}\).</p>]]></text></generalfeedback>
  <defaultgrade>1</defaultgrade><penalty>0.1</penalty><hidden>0</hidden><idnumber></idnumber>
  <stackversion><text>2025073100</text></stackversion>
  <questionvariables><text>C1: rand(4)+1;
C2: rand(4)+1;
L: 3;
u0: C2;
uL: C1*L + C2;
ta1: C1*x + C2;</text></questionvariables>
  <specificfeedback format="html"><text></text></specificfeedback>
  <questionnote format="html"><text><![CDATA[<p>u(x)={@ta1@}</p>]]></text></questionnote>
  <questiondescription format="html"><text></text></questiondescription>
  <questionsimplify>1</questionsimplify><assumepositive>0</assumepositive><assumereal>0</assumereal>
  <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
  <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
  <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
  <decimals>.</decimals><scientificnotation>*10</scientificnotation><multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign><complexno>i</complexno><inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol><matrixparens>[</matrixparens><isbroken>0</isbroken><variantsselectionseed></variantsselectionseed>
  <input>
    <name>ans1</name><type>algebraic</type><tans>ta1</tans><boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax><insertstars>0</insertstars><syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute><forbidwords></forbidwords><allowwords></allowwords>
    <forbidfloat>0</forbidfloat><requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype><mustverify>1</mustverify><showvalidation>3</showvalidation><options></options>
  </input>
  <prt>
    <name>prt1</name><value>1.0000000</value><autosimplify>1</autosimplify><feedbackstyle>1</feedbackstyle>
    <feedbackvariables><text></text></feedbackvariables>
    <node>
      <name>0</name><description></description><answertest>AlgEquiv</answertest>
      <sans>ans1</sans><tans>ta1</tans><testoptions></testoptions><quiet>0</quiet>
      <truescoremode>=</truescoremode><truescore>1</truescore><truepenalty></truepenalty>
      <truenextnode>-1</truenextnode><trueanswernote>prt1-1-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode><falsescore>0</falsescore><falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode><falseanswernote>prt1-1-F</falseanswernote>
      <falsefeedback format="html"><text></text></falsefeedback>
    </node>
  </prt>
</question>

<question type="stack">
  <name><text>PDE BVP - Q2: u_xx=0, u(a)=p, u(b)=q</text></name>
  <questiontext format="html"><text><![CDATA[<p>Solve the boundary value problem:</p>
<p>\[\frac{\partial^2 u}{\partial x^2} = 0, \quad u({@a@}) = {@ua@}, \quad u({@b@}) = {@ub@}.\]</p>
<p>\(u(x) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text></questiontext>
  <generalfeedback format="html"><text><![CDATA[<h3>Solution</h3>
<p>General solution: \(u(x) = C_1 x + C_2\).</p>
<p>BC1: \({@a@}C_1 + C_2 = {@ua@}\).</p>
<p>BC2: \({@b@}C_1 + C_2 = {@ub@}\).</p>
<p>Subtracting: \(({@b-a@})C_1 = {@ub-ua@}\), so \(C_1 = {@C1@}\), \(C_2 = {@C2@}\).</p>
<p>Therefore \(u(x) = {@ta1@}\).</p>]]></text></generalfeedback>
  <defaultgrade>1</defaultgrade><penalty>0.1</penalty><hidden>0</hidden><idnumber></idnumber>
  <stackversion><text>2025073100</text></stackversion>
  <questionvariables><text>C1: rand(4)+1;
C2: rand(4)+1;
a: 1;
b: rand(3)+3;
ua: C1*a + C2;
ub: C1*b + C2;
ta1: C1*x + C2;</text></questionvariables>
  <specificfeedback format="html"><text></text></specificfeedback>
  <questionnote format="html"><text><![CDATA[<p>u(x)={@ta1@}</p>]]></text></questionnote>
  <questiondescription format="html"><text></text></questiondescription>
  <questionsimplify>1</questionsimplify><assumepositive>0</assumepositive><assumereal>0</assumereal>
  <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
  <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
  <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
  <decimals>.</decimals><scientificnotation>*10</scientificnotation><multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign><complexno>i</complexno><inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol><matrixparens>[</matrixparens><isbroken>0</isbroken><variantsselectionseed></variantsselectionseed>
  <input>
    <name>ans1</name><type>algebraic</type><tans>ta1</tans><boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax><insertstars>0</insertstars><syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute><forbidwords></forbidwords><allowwords></allowwords>
    <forbidfloat>0</forbidfloat><requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype><mustverify>1</mustverify><showvalidation>3</showvalidation><options></options>
  </input>
  <prt>
    <name>prt1</name><value>1.0000000</value><autosimplify>1</autosimplify><feedbackstyle>1</feedbackstyle>
    <feedbackvariables><text></text></feedbackvariables>
    <node>
      <name>0</name><description></description><answertest>AlgEquiv</answertest>
      <sans>ans1</sans><tans>ta1</tans><testoptions></testoptions><quiet>0</quiet>
      <truescoremode>=</truescoremode><truescore>1</truescore><truepenalty></truepenalty>
      <truenextnode>-1</truenextnode><trueanswernote>prt1-1-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode><falsescore>0</falsescore><falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode><falseanswernote>prt1-1-F</falseanswernote>
      <falsefeedback format="html"><text></text></falsefeedback>
    </node>
  </prt>
</question>

<question type="stack">
  <name><text>PDE BVP - Q3: IVP u_t=0</text></name>
  <questiontext format="html"><text><![CDATA[<p>Solve the initial value problem:</p>
<p>\[\frac{\partial u}{\partial t} = 0, \quad u(x,0) = {@ic_label@}.\]</p>
<p>\(u(x,t) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text></questiontext>
  <generalfeedback format="html"><text><![CDATA[<h3>Solution</h3>
<p>Since \(\frac{\partial u}{\partial t}=0\), \(u\) is independent of \(t\), so \(u(x,t)=f(x)\).</p>
<p>From \(u(x,0)={@ic_label@}\): \(f(x)={@ic_label@}\).</p>
<p>Therefore \(u(x,t)={@ta1@}\).</p>]]></text></generalfeedback>
  <defaultgrade>1</defaultgrade><penalty>0.1</penalty><hidden>0</hidden><idnumber></idnumber>
  <stackversion><text>2025073100</text></stackversion>
  <questionvariables><text>ic_idx: rand(4);
ic_list: [cos(x), sin(x), x^2+1, exp(x)];
ic_label_list: ["\\cos x", "\\sin x", "x^2+1", "e^x"];
ta1: ic_list[ic_idx+1];
ic_label: ic_list[ic_idx+1];</text></questionvariables>
  <specificfeedback format="html"><text></text></specificfeedback>
  <questionnote format="html"><text><![CDATA[<p>u(x,t)={@ta1@}</p>]]></text></questionnote>
  <questiondescription format="html"><text></text></questiondescription>
  <questionsimplify>1</questionsimplify><assumepositive>0</assumepositive><assumereal>0</assumereal>
  <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
  <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
  <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
  <decimals>.</decimals><scientificnotation>*10</scientificnotation><multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign><complexno>i</complexno><inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol><matrixparens>[</matrixparens><isbroken>0</isbroken><variantsselectionseed></variantsselectionseed>
  <input>
    <name>ans1</name><type>algebraic</type><tans>ta1</tans><boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax><insertstars>0</insertstars><syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute><forbidwords></forbidwords><allowwords></allowwords>
    <forbidfloat>0</forbidfloat><requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype><mustverify>1</mustverify><showvalidation>3</showvalidation><options></options>
  </input>
  <prt>
    <name>prt1</name><value>1.0000000</value><autosimplify>1</autosimplify><feedbackstyle>1</feedbackstyle>
    <feedbackvariables><text></text></feedbackvariables>
    <node>
      <name>0</name><description></description><answertest>AlgEquiv</answertest>
      <sans>ans1</sans><tans>ta1</tans><testoptions></testoptions><quiet>0</quiet>
      <truescoremode>=</truescoremode><truescore>1</truescore><truepenalty></truepenalty>
      <truenextnode>-1</truenextnode><trueanswernote>prt1-1-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode><falsescore>0</falsescore><falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode><falseanswernote>prt1-1-F</falseanswernote>
      <falsefeedback format="html"><text></text></falsefeedback>
    </node>
  </prt>
</question>

<question type="stack">
  <name><text>PDE BVP - Q4: Classify Dirichlet/Neumann/Mixed</text></name>
  <questiontext format="html"><text><![CDATA[<p>State whether the following boundary conditions are <strong>Dirichlet</strong>, <strong>Neumann</strong>, or <strong>Mixed (Robin)</strong>:</p>
<p>(a) \(u(0,t) = 0\): [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>(b) \(\dfrac{\partial u}{\partial x}(1,t) = 5\): [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
<p>(c) \(u(0,t) + \dfrac{\partial u}{\partial x}(0,t) = 1\): [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>]]></text></questiontext>
  <generalfeedback format="html"><text><![CDATA[<h3>Solution</h3>
<p>(a) \(u(0,t)=0\) specifies the function value &rarr; <strong>Dirichlet</strong>.</p>
<p>(b) \(u_x(1,t)=5\) specifies the derivative &rarr; <strong>Neumann</strong>.</p>
<p>(c) Involves both function and derivative &rarr; <strong>Mixed (Robin)</strong>.</p>]]></text></generalfeedback>
  <defaultgrade>3</defaultgrade><penalty>0.1</penalty><hidden>0</hidden><idnumber></idnumber>
  <stackversion><text>2025073100</text></stackversion>
  <questionvariables><text>ta1: [[1, true, "Dirichlet"], [2, false, "Neumann"], [3, false, "Mixed (Robin)"]];
ta2: [[1, false, "Dirichlet"], [2, true, "Neumann"], [3, false, "Mixed (Robin)"]];
ta3: [[1, false, "Dirichlet"], [2, false, "Neumann"], [3, true, "Mixed (Robin)"]];</text></questionvariables>
  <specificfeedback format="html"><text></text></specificfeedback>
  <questionnote format="html"><text><![CDATA[<p>(a) Dirichlet, (b) Neumann, (c) Mixed.</p>]]></text></questionnote>
  <questiondescription format="html"><text></text></questiondescription>
  <questionsimplify>1</questionsimplify><assumepositive>0</assumepositive><assumereal>0</assumereal>
  <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
  <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
  <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
  <decimals>.</decimals><scientificnotation>*10</scientificnotation><multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign><complexno>i</complexno><inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol><matrixparens>[</matrixparens><isbroken>0</isbroken><variantsselectionseed></variantsselectionseed>
  <input>
    <name>ans1</name><type>radio</type><tans>ta1</tans><boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax><insertstars>0</insertstars><syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute><forbidwords></forbidwords><allowwords></allowwords>
    <forbidfloat>0</forbidfloat><requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype><mustverify>0</mustverify><showvalidation>0</showvalidation><options></options>
  </input>
  <input>
    <name>ans2</name><type>radio</type><tans>ta2</tans><boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax><insertstars>0</insertstars><syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute><forbidwords></forbidwords><allowwords></allowwords>
    <forbidfloat>0</forbidfloat><requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype><mustverify>0</mustverify><showvalidation>0</showvalidation><options></options>
  </input>
  <input>
    <name>ans3</name><type>radio</type><tans>ta3</tans><boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax><insertstars>0</insertstars><syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute><forbidwords></forbidwords><allowwords></allowwords>
    <forbidfloat>0</forbidfloat><requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype><mustverify>0</mustverify><showvalidation>0</showvalidation><options></options>
  </input>
  <prt>
    <name>prt1</name><value>1.0000000</value><autosimplify>1</autosimplify><feedbackstyle>1</feedbackstyle>
    <feedbackvariables><text></text></feedbackvariables>
    <node>
      <name>0</name><description></description><answertest>AlgEquiv</answertest>
      <sans>ans1</sans><tans>mcq_correct(ta1)</tans><testoptions></testoptions><quiet>0</quiet>
      <truescoremode>=</truescoremode><truescore>1</truescore><truepenalty></truepenalty>
      <truenextnode>-1</truenextnode><trueanswernote>prt1-1-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode><falsescore>0</falsescore><falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode><falseanswernote>prt1-1-F</falseanswernote>
      <falsefeedback format="html"><text></text></falsefeedback>
    </node>
  </prt>
  <prt>
    <name>prt2</name><value>1.0000000</value><autosimplify>1</autosimplify><feedbackstyle>1</feedbackstyle>
    <feedbackvariables><text></text></feedbackvariables>
    <node>
      <name>0</name><description></description><answertest>AlgEquiv</answertest>
      <sans>ans2</sans><tans>mcq_correct(ta2)</tans><testoptions></testoptions><quiet>0</quiet>
      <truescoremode>=</truescoremode><truescore>1</truescore><truepenalty></truepenalty>
      <truenextnode>-1</truenextnode><trueanswernote>prt2-1-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode><falsescore>0</falsescore><falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode><falseanswernote>prt2-1-F</falseanswernote>
      <falsefeedback format="html"><text></text></falsefeedback>
    </node>
  </prt>
  <prt>
    <name>prt3</name><value>1.0000000</value><autosimplify>1</autosimplify><feedbackstyle>1</feedbackstyle>
    <feedbackvariables><text></text></feedbackvariables>
    <node>
      <name>0</name><description></description><answertest>AlgEquiv</answertest>
      <sans>ans3</sans><tans>mcq_correct(ta3)</tans><testoptions></testoptions><quiet>0</quiet>
      <truescoremode>=</truescoremode><truescore>1</truescore><truepenalty></truepenalty>
      <truenextnode>-1</truenextnode><trueanswernote>prt3-1-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode><falsescore>0</falsescore><falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode><falseanswernote>prt3-1-F</falseanswernote>
      <falsefeedback format="html"><text></text></falsefeedback>
    </node>
  </prt>
</question>

<question type="stack">
  <name><text>PDE BVP - Q5: u_x=0, uniqueness and solution</text></name>
  <questiontext format="html"><text><![CDATA[<p>The PDE \(\dfrac{\partial u}{\partial x} = 0\) has infinitely many solutions without an extra condition.</p>
<p>Given the condition \(u(0,y) = {@ic@}\), find the unique solution.</p>
<p>\(u(x,y) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text></questiontext>
  <generalfeedback format="html"><text><![CDATA[<h3>Solution</h3>
<p>Integrating \(u_x=0\) w.r.t. \(x\) gives \(u(x,y) = f(y)\), an arbitrary function of \(y\).</p>
<p>Applying \(u(0,y) = {@ic@}\): \(f(y) = {@ic@}\).</p>
<p>Therefore \(u(x,y) = {@ta1@}\).</p>]]></text></generalfeedback>
  <defaultgrade>1</defaultgrade><penalty>0.1</penalty><hidden>0</hidden><idnumber></idnumber>
  <stackversion><text>2025073100</text></stackversion>
  <questionvariables><text>ic_idx: rand(4);
ic_list: [y^2, sin(y), y^3+1, exp(y)];
ic: ic_list[ic_idx+1];
ta1: ic;</text></questionvariables>
  <specificfeedback format="html"><text></text></specificfeedback>
  <questionnote format="html"><text><![CDATA[<p>u(x,y)={@ta1@}</p>]]></text></questionnote>
  <questiondescription format="html"><text></text></questiondescription>
  <questionsimplify>1</questionsimplify><assumepositive>0</assumepositive><assumereal>0</assumereal>
  <prtcorrect format="html"><text><![CDATA[<span style="color:green;">Correct!</span>]]></text></prtcorrect>
  <prtpartiallycorrect format="html"><text><![CDATA[<span style="color:orange;">Partially correct.</span>]]></text></prtpartiallycorrect>
  <prtincorrect format="html"><text><![CDATA[<span style="color:red;">Incorrect.</span>]]></text></prtincorrect>
  <decimals>.</decimals><scientificnotation>*10</scientificnotation><multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign><complexno>i</complexno><inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol><matrixparens>[</matrixparens><isbroken>0</isbroken><variantsselectionseed></variantsselectionseed>
  <input>
    <name>ans1</name><type>algebraic</type><tans>ta1</tans><boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax><insertstars>0</insertstars><syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute><forbidwords></forbidwords><allowwords></allowwords>
    <forbidfloat>0</forbidfloat><requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype><mustverify>1</mustverify><showvalidation>3</showvalidation><options></options>
  </input>
  <prt>
    <name>prt1</name><value>1.0000000</value><autosimplify>1</autosimplify><feedbackstyle>1</feedbackstyle>
    <feedbackvariables><text></text></feedbackvariables>
    <node>
      <name>0</name><description></description><answertest>AlgEquiv</answertest>
      <sans>ans1</sans><tans>ta1</tans><testoptions></testoptions><quiet>0</quiet>
      <truescoremode>=</truescoremode><truescore>1</truescore><truepenalty></truepenalty>
      <truenextnode>-1</truenextnode><trueanswernote>prt1-1-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode><falsescore>0</falsescore><falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode><falseanswernote>prt1-1-F</falseanswernote>
      <falsefeedback format="html"><text></text></falsefeedback>
    </node>
  </prt>
</question>

</quiz>
